# Week 3 Problem Set

## Homeworks

**HW1.** *Fibonaci:* Write a function to find the Fibonacci number given an index. These are the example of a the first few numbers in Fibonacci series.

| Indices    | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7  | 8  | 9  |
|------------|---|---|---|---|---|---|---|----|----|----|
| The series | 0 | 1 | 1 | 2 | 3 | 5 | 8 | 13 | 21 | 34 |

We can write that the Fibonacci number at index *i* is given by:

$fibo(i) = fibo(i-1) + fibo(i-2)$

Use recursion for your implementation.

Reorder the sequence of the codes to implement the right solution. 

```python
codes = {
    'e': 'def fibonacci(index):',
        'h': 'if index == 0:',
            'a': 'return 0',
        'g': 'elif index == 1:',
            'b': 'return 1',
        'c': 'return fibonacci(index - 1) + fibonacci(index - 2)',
    # 'd': 'return fibonacci(index) + fibonacci(index - 1)',
    # 'f': 'else',
}
```

```python
# Enter your sequence of code and use "sub" or "exit sub" as necessary
answer = ['e', 'sub', 'h', 'sub', 'a', 'exit sub', 'g', 'sub', 'b', 'exit sub', 'c']
###
### YOUR CODE HERE
###

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Write and test the code below.

```python
def fibonacci(index):
    if index == 0:
        return 0
    elif index == 1:
        return 1
    return fibonacci(index - 1) + fibonacci(index - 2)
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(3) == 2
assert fibonacci(7) == 13
assert fibonacci(9) == 34
```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Max-Heapify:* Recall the algorithm for max-heapify in restoring the *max-heap property* of a binary heap tree. In the previous implementation, we used iteration. Implement the same algorithm using recursion. 

Reorder the sequence of the code to implement max heapify function.

```python
codes = {
	'd': 'def max_heapify(array, index, heap_size):',
		'a': 'largest = index',
		'e': '''left_idx = left_of(index)
				right_idx = right_of(index)''',
		'i': 'if (left_idx < heap_size) and (array[left_idx] > array[index]):',
			'b': 'largest = left_idx',
		'j': 'if (right_idx < heap_size) and (array[right_idx] > array[largest]):',
			'c': 'largest = right_idx',
		'h': 'if largest != index:',
			'k': 'array[index], array[largest] = array[largest], array[index]',
			'g': 'max_heapify(array, largest, heap_size)',
		'f': 'else:',
}
```

```python
# enter your code sequence below and use "sub" or "exit sub" as necessary.
answer = ['d', 'sub', None, None, 'sub', 'b', 'exit sub', 'f', 'sub', None, 'exit sub', None, 'sub', None, 'exit sub', 'h', 'sub', 'k', None]
###
### YOUR CODE HERE
###

```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def left_of(index):
    return (2 * index) + 1

def right_of(index):
    return left_of(index) + 1

def max_heapify(array, index, size):
    left, right = left_of(index), right_of(index)
    is_leaf = left >= size
    is_left_only = right >=  size
    if is_leaf or (array[index] >= array[left] and (is_left_only or array[index] >= array[right])):
        return
    
    if is_left_only or array[left] > array[right]:
        array[left], array[index] = array[index], array[left]
        max_heapify(array, left, size)
    else:
        array[right], array[index] = array[index], array[right]
        max_heapify(array, right, size)
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

<!-- grade_id: hw23, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *String Permutation:* Write two functions to return an array of all the permutations of a string. For example, if the input is `abc`, the output should be

```
output = ["abc", "acb", "bac", "bca", "cab", "cba"]
```

The first function only has one argument which is the input string, i.e. `permutate(s)`. The second function is the recursive function with two arguments String 1 (`str1`) and and String 2 (`str2`). The first function calls the second method `permutate("", s)` at the beginning. The second function should use a loop to move a character from `str2` to `str1` and recursively invokes it with a new `str1` and `str2`. 

```python
def permutate(s):
    if len(s) == 1:
        return [s]
    res = []
    for i, char in enumerate(s):
        remaining_string = "".join(x for remaining_i, x in enumerate(s) if remaining_i != i)
        permutations = permutate(remaining_string)
        res += [char + permutation for permutation in permutations]
    # print(res)
    return res
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert set(permutate("abc"))==set(["abc", "acb", "bac", "bca", "cab", "cba"])
assert set(permutate("abcd"))==set(["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"])
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *GCD:* Implement the recursive implementation of Euclid's algorithm for Greatest Common Divisor (GCD).

Hint: [Wikipedia: Greatest Common Divisor](https://en.wikipedia.org/wiki/Greatest_common_divisor#Euclid's_algorithm)

```python
def gcd(a, b):
    if b == 0:
        return a
    return gcd(b, a % b)
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert gcd(0,0) == 0
assert gcd(2,0) == 2
assert gcd(0,2) == 2
assert gcd(2,3) == 1
assert gcd(4, 28) == 4
assert gcd(1024, 526) == 2
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```