```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 1 Problem Set

## Cohort Problems

**CS0.** Implement factorial problem using iteration. The function should takes in an Integer input and returns and Integer output which is the factorial of the input. Recall that:

$n! = n \times (n-1) \times (n-2) \times \ldots \times 2 \times 1$

You should consider the case when $n$ is zero and one as well. 

```python
# iteration
def factorial_iteration(n):
    result = 1
    for i in range(n, 1, -1):
        result *= i
    return result
```

<!-- grade_id: cs01, points: 1, grade: true, solution: false -->
```python
assert factorial_iteration(0) == 1
assert factorial_iteration(1) == 1
assert factorial_iteration(5) == 120
assert factorial_iteration(7) == 5040
assert factorial_iteration(11) == 39916800
```

<!-- grade_id: cs02, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS1-a.** Enter the steps of the following pseudocode as a list to implement Bubble Sort version 1. Enter your answer by creating a list with the right sequence. Add "sub" if the subsequent steps are steps "under" (sub) the previous step. Use "exit sub" before steps that is not under the overarching steps. See example below for the first few steps of Bubble Sort pseudocode and enter the remaining steps. 

```
1. n = length of array
2. For outer_index from 1 to n-1, do:
    2.1 For inner_index from 1 to n-1, do:
        2.1.1 first_number = array[inner_index - 1]
        2.1.2 second_number = array[inner_index]
        2.1.3 if first_number > second_number, do:
            2.1.3.1 swap(first_number, second_number)
```

Reference: [Bubble Sort and Insertion Sort](https://data-driven-world.github.io/notes/bubble_insertion_sort)  



```python
steps = {'1': 'n = length of array',
         '2': 'For outer_index from 1 to n-1, do:',
         '2.1': 'For inner_index from 1 to n-1, do:',
         '2.1.1': 'first_number = array[inner_index - 1]',
         '2.1.2': 'second_number = array[inner_index]',
         '2.1.3': 'if first_number > second_number, do:',
         '2.1.3.1': 'swap(first_number, second_number)}'}

# replace the None with the right key from the steps variable above or "sub"/"exit sub"
answer = ['1', '2', 'sub', '2.1', 'sub', '2.1.1', '2.1.2', '2.1.3', 'sub', '2.1.3.1']
###
### YOUR CODE HERE
###

```

<!-- grade_id: cs1a1, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```


**CS1-b.** Create a function that implements **Bubble Sort version 1** (from your Notes) to sort an array of integers. The function should sort the input array in place. Refer to the above pseudocode.

```python
def bubble_sort(array):
    n = len(array)
    for outer_index in range(1, n):
        for inner_index in range(1, n):
            first_number = array[inner_index - 1]
            second_number = array[inner_index]
            if first_number > second_number:
                array[inner_index - 1], array[inner_index] = array[inner_index], array[inner_index - 1]
```

<!-- grade_id: cs1b1, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
bubble_sort(array)
print(array)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
```

<!-- grade_id: cs1b2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Modify CS1, so that it returns  the number of comparisons that are performed.

Hint: To count the number of comparisons, you can create an integer variable which you *increment* just before the `if` statement. 

```python
def bubble_sort(array):
    count = 0
    n = len(array)
    for outer_index in range(1, n):
        for inner_index in range(1, n):
            first_number = array[inner_index - 1]
            second_number = array[inner_index]
            count += 1
            if first_number > second_number:
                array[inner_index - 1], array[inner_index] = array[inner_index], array[inner_index - 1]
    return count
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 121
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 121
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3-a.** Sequence the steps below to have a correct pseudocode for Insertion sort. 

```python
steps = {'a': 'For outer_index in Range(from 1 to n-1), do:',
         'b': 'As long as (inner_index > 0) AND (array[inner_index] < array[inner_index - 1]), do:',
         'c': 'inner_index = inner_index - 1',
         'd': 'n = length of array',
         'e': 'inner_index = outer_index',
         'f': 'swap(array[inner_index - 1], array[inner_index])'}

# add "sub" or "exit sub" as necessary
answer = ['d', 'a', 'sub', 'e', 'b', 'sub', 'f', 'c']

###
### YOUR CODE HERE
###

```

<!-- grade_id: cs3a1, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3-b.** Create a function that implements **Insertion Sort** to sort an array of floats. The function should sort the input array in place. 

```python
def insertion_sort(array):
    n = len(array)
    for i in range(1, n):
        # while array[i - 1] > array[i] and i >= 1:
        #     array[i - 1], array[i] = array[i], array[i - 1]
        #     i -= 1
        current_val = array[i]
        while array[i - 1] > current_val and i >= 1:
            array[i] = array[i - 1]
            i -= 1
        array[i] = current_val

# insertion_sort = list.sort
```

<!-- grade_id: cs3b1, points: 1, grade: true, solution: false -->
```python
array = [10.3,3.8,8.4,47.1,1.0,0,-39.8,8.4,4.7,7.6,-6.5,-5.0]
insertion_sort(array)
print(array)
assert array == [-39.8, -6.5, -5.0, 0.0, 1.0, 3.8, 4.7, 7.6, 8.4, 8.4, 10.3, 47.1]
```

Output:
```
[-39.8, -6.5, -5.0, 0, 1.0, 3.8, 4.7, 7.6, 8.4, 8.4, 10.3, 47.1]
```

<!-- grade_id: cs3b2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Write a function `gen_random_int` that generates a list of integers from `0` to `n - 1` with its sequence randomly shuffled. The function should take in an integer `n`, denoting the number of integers to be generated. 

Hint: 
- You can use `random.shuffle(mylist)`  to shuffle the elements of a list called `mylist`.
- Refer to [Python's Random module](https://docs.python.org/3.7/library/random.html) for more details on how to use `shuffle()`.

```python
import random

def gen_random_int(number, seed):
    random.seed(seed)
    res = list(range(number))
    random.shuffle(res)
    return res
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
output = gen_random_int(10, 100)
print(output)
assert output == [4, 0, 5, 9, 3, 1, 6, 8, 7, 2]
```

Output:
```
[4, 0, 5, 9, 3, 1, 6, 8, 7, 2]
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```