```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 1 Problem Set
## Homework



**HW0-a**. Reorder the code to check if a string is a palindrome. 

```python
steps = { 'a': 'return False',
          'b': 'return True',
          'c': 'if s[idx] != s[len(s)-idx-1]:',
          'd': 'for idx in range(len(s)//2):',
          'e': 'def palindrome(s):'
        }

# enter "sub" or "exit sub" as necessary
answer = ['e', 'sub', 'd', 'sub', 'c', 'sub', 'a', 'end sub', 'b']

###
### YOUR CODE HERE
###

```

<!-- grade_id: hw0a1, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW0-b.** Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). Use iteration for this solution.


```python
def palindrome(s):
    # sir please why not just this
    return s == s[::-1]
    # iteration (for no reason):
    # for idx in range(len(s)//2):
    #     if s[idx] != s[len(s)-idx-1]:
    #         return False
    # return True
```

<!-- grade_id: hw0b1, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 
```

<!-- grade_id: hw0b2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW-1-a.** Reorder the correct sequence of code to implement bubble sort algorithm version 2. You can check the pseudocode in [Bubble Sort and Insertion Sort](https://data-driven-world.github.io/2023/notes/sorting-algorithms/bubble-sort). 

The function should ...

- **modify/mutate** an existing array into a sorted one
- **return** the number of comparisons made

Hint: To count the number of comparisons made you can simply increment an integer variable (counter) right before the `if` statement.

```python
codes = {
        'i': 'def bubble_sort(array):',
        'b': '''count = 0
                n = len(array)
                swapped = True
        ''',
        'e': 'while swapped:',
        'a': 'swapped = False',
        'h': 'for idx in range(1, n):',
        'c': 'count += 1',
        'g': 'if array[idx-1] > array[idx]:',
        'f': 'array[idx - 1], array[idx] = array[idx], array[idx - 1]',
        'j': 'swapped = True',
        'd': 'return count',
}

# enter "sub" or "exit sub" as necessary
# you can also use code more than once
answer = ['i', 'b', 'e', 'a', 'sub', 'h', 'c', 'g', 'sub', 'f', 'j', 'end sub', 'end sub', 'd']

###
### YOUR CODE HERE
###

```

<!-- grade_id: hw1a1, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW1-b.** Create a function to implement **Bubble Sort Algorithm version 2**. 

```python
def bubble_sort(array):
    count = 0
    n = len(array)
    swapped = True
    while swapped:
        swapped = False
        for i in range(0, n - 1): # up to second-last item
            count += 1
            if array[i] > array[i + 1]:
                array[i], array[i + 1] = array[i + 1], array[i]
                swapped = True
    return count
```

<!-- grade_id: hw1b1, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 121
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 121
```

<!-- grade_id: hw1b2, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** The solution in HW1 can be improved!

The n-th pass places the n-th largest element into its final place (i.e. after 1 cycle, the (1st) largest element will be in its correct position).

Hence, we can reduce subsequent iterations by only considering the other (unsorted) `n-1` elements in the array. 

Implement the algorithm of **Bubble Sort version 4** under [Bubble Sort and Insertion Sort](https://data-driven-world.github.io/2023/notes/sorting-algorithms/bubble-sort). 


```python
def bubble_sort(array):
    count = 0
    n = len(array)
    swapped = True
    while swapped:
        swapped = False
        new_n = 0
        for i in range(0, n - 1): # up to second-last item
            count += 1
            if array[i] > array[i + 1]:
                array[i], array[i + 1] = array[i + 1], array[i]
                swapped = True
                new_n = i + 1
        n = new_n
    return count
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 66
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 66
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```