# Week 5 Problem Set

## Cohort Sessions

**CS1.** *Dictionary:* Implement a Graph using a *Dictionary* where the keys are the Vertices in the Graph and the values (in the the key-value pair) correspond to an Array containing the neighbouring Vertices. For example, let's represent the following graph:
```
    A -> B
    A -> C
    B -> C
    B -> D
    C -> D
    D -> C
    E -> F
    F -> C
```
Create a Dictionary to represent the graph above.

```python
# replace the None with a dictionary representing the graph
graph = {
    'A': ['B', 'C'],
    'B': ['C', 'D'],
    'C': ['D'],
    'D': ['C'],
    'E': ['F'],
    'F': ['C']
}

###
### YOUR CODE HERE
###

```

```python
print(graph)
```

Output:
```
{'A': ['B', 'C'], 'B': ['C', 'D'], 'C': ['D'], 'D': ['C'], 'E': ['F'], 'F': ['C']}
```

Write a function `get_neighbours(graph, vert)` which returns a list of all neighbours of the requested Vertex `vert` in the `graph`. Return `None` if the Vertex is not in the graph.

```python
def get_neighbours(graph: dict, target_vertex):
    res = []
    for vertex, edges in graph.items():
        if vertex == target_vertex:
            res += edges
        # elif target_vertex in edges:
        #     res.append(vertex)
    return res if len(res) != 0 else None
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
print(get_neighbours(graph, "B"))
assert get_neighbours(graph, "B") == ["C", "D"]
assert get_neighbours(graph, "A") == ["B", "C"]
assert get_neighbours(graph, "F") == ["C"]
assert get_neighbours(graph, "Z") == None
```

Output:
```
['C', 'D']
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Write a function `get_source(graph, vert)` which returns a list of all source Vertices pointing to `vert` in the `graph`. For example, Vertex "C" has the following source Vertices: `["A", "B", "D", "F"]`. Return a `None` object if there are none.

```python
def get_source(graph, target_vert):
    res = []
    for vertex, edges in graph.items():
        if target_vert in edges:
            res.append(vertex)
    
    return None if len(res) == 0 else res
```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
assert sorted(get_source(graph, "C")) == ["A", "B", "D", "F"]
assert sorted(get_source(graph, "D")) == ["B", "C"]
assert sorted(get_source(graph, "F")) == ["E"]
assert get_source(graph, "Z") == None

```

<!-- grade_id: cs14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Create a class `Vertex` to represent a vertex in a graph. The class `Vertex` has the following attributes:
- `id_`: to identify each vertex. This is of String data type.
- `neighbours`: which is a Dictionary where the keys are the neighbouring `Vertex` object instances that are connected to the current Vertex and the values are the weights of the edge between the current Vertex and the neighbouring vertices. 

The class should also have the following methods:

- `__init__(self, id_)`: which is used to initialized the attribute `id_`. By default, `id_` is set to an empty String . The attribute `neighbours` is always set to an empty dictionary.
- `add_neighbour(self, nbr_vertex, weight)`: which adds a neighbouring Vertex to the current Vertex. The second argument provides the weight of the edge between the current Vertex and the newly added neighbouring Vertex. By default, `weight` is `0`.
- `get_neigbours(self)`: which returns all the Vertices connected to the current Vertex as a list. The elements of the output list are of `Vertex` object instances.
- `get_weight(self, neighbour)`: which returns the weight of the requested neighbour. It should return `None` if the requested neighbour is not found.
- `__eq__(self, other)`: which returns true if the id of the current vertex object is the same as the `other` vertex's id. 
- `__lt__(self, other)`: which returns true if the id of the current vertex object is less than the `other` vertex's id.
- `__hash__(self)`: which calls the `hash()` function on `id_` and returns it. This allows the object to be a dictionary key. This is provided for you.
- `__str__(self)`: This method should return the id of the current vertex and a list of `id_`s of the neighbouring vertices, like `Vertex 2 is connected to: 3, 4, 5` .

```python
class Vertex:
    def __init__(self, id_=""):
        self.id_ = id_
        self.neighbours: dict[Vertex, int] = {}
    
    def add_neighbour(self, nbr_vertex, weight: int = 0):
        self.neighbours[nbr_vertex] = weight
    
    def get_neighbours(self):
        return list(self.neighbours.keys())
    
    def get_weight(self, neighbour):
        return self.neighbours.get(neighbour)
    
    def __eq__(self, other):
        return self.id_ == other.id_
    
    def __lt__(self, other):
        return self.id_ < other.id_
    
    def __hash__(self):
        return hash(self.id_)
    
    def __str__(self):
        return f'Vertex {self.id_} is connected to: {", ".join(neighbour.id_ for neighbour in self.get_neighbours())}'
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
v1 = Vertex("1")
assert v1.id_ == "1" and len(v1.neighbours) == 0
v2 = Vertex("2")
v1.add_neighbour(v2)
assert v1.get_neighbours()[0].id_ == "2" and v1.neighbours[v1.get_neighbours()[0]] == 0
v3 = Vertex("3")
v1.add_neighbour(v3, 3)
assert v1.get_weight(v3) == 3
v4 = Vertex("4")
assert v1.get_weight(v4) == None
assert v1 < v2
assert v1 != v2
assert str(v1) == "Vertex 1 is connected to: 2, 3"
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** Create a class `Graph` to represent a Graph. The class has the following attribute:
- `vertices`: which is a *dictionary* of Vertices. The keys are the `id`s of the Vertices and the values are `Vertex` object instances.

The class has the following property:
- `num_vertices`: which is a *computed* property that returns the number of vertices in the graph.

The class also has the following methods:
- `__init__(self)`: which initializes the graph with an empty dictionary.
- `_create_vertex(self, id_)`: which creates a new `Vertex` object with a given `id_`. This method is never called directly and is only used by `add_vertex(id_)`.
- `add_vertex(self, id_)`: which creates a new `Vertex` object, adding it into the dictionary `vertices`. The argument `id_` is a String. This method should call `_create_vertex(id_)`.
- `get_vertex(self, id_)`: which returns the `Vertex` object instance of the requested `id_`. The method should return `None` if the requested `id_` cannot be found. The argument `id_` is a String.
- `add_edge(start_v, end_v)`: which creates an edge from one Vertex to another Vertex. The arguments are the `id_`s of the two vertices and are both Strings.
- `get_neighbours(self, id_)`: which returns a list of `id_`s all the neighbouring vertices (of the specified Vertex `id_`). It should return `None` if `id_` cannot be found. The argument `id_` is a String and the elements of the output list are of `str` data type. 
- `__contains__(self, id_)`: which returns either `True` or `False` depending on whether the graph contains the specified Vertex's `id_`. The argument `id_` is a String.

```python
class Graph:
    def __init__(self):
        self.vertices: dict[str, Vertex] = {}
        
    def _create_vertex(self, id_):
        return Vertex(id_)
    
    def add_vertex(self, id_):
        self.vertices[id_] = self._create_vertex(id_)
    
    def get_vertex(self, id_):
        return self.vertices.get(id_)
    
    def add_edge(self, start_v, end_v, weight=0):
        if self.get_vertex(start_v) is None:
            raise ValueError('start_v was not in graph')
        elif self.get_vertex(end_v) is None:
            raise ValueError('end_v was not in graph')

        self.get_vertex(start_v).add_neighbour(self.get_vertex(end_v), weight)

    def get_neighbours(self, id_):
        if self.get_vertex(id_) is None:
            return None
        return [neighbour.id_ for neighbour in self.get_vertex(id_).get_neighbours()]
    
    def __contains__(self, val):
        return self.get_vertex(val) is not None
    
    def __iter__(self):
        for k,v in self.vertices.items():
            yield v 
        
    # write a code to create a computed property called num_vertices
    @property
    def num_vertices(self):
        return len(self.vertices)

```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
g = Graph()
assert g.vertices == {} and g.num_vertices == 0
g.add_vertex("A")
g.add_vertex("B")
g.add_vertex("C")
g.add_vertex("D")
g.add_vertex("E")
g.add_vertex("F")
assert g.num_vertices == 6
assert "A" in g
assert "B" in g
assert "C" in g
assert "D" in g
assert "E" in g
assert "F" in g
g.add_edge("A", "B")
g.add_edge("A", "C")
g.add_edge("B", "C")
g.add_edge("B", "D")
g.add_edge("C", "D")
g.add_edge("D", "C")
g.add_edge("E", "F")
g.add_edge("F", "C")
assert sorted(g.get_neighbours("A")) == ["B", "C"]
assert sorted(g.get_neighbours("B")) == ["C", "D"]
assert sorted(g.get_neighbours("C")) == ["D"]
assert [v.id_ for v in g] == ["A", "B", "C", "D", "E", "F"]
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Create a subclass of `Vertex` called `VertexSearch`. This class has the following additional attributes:
- `colour`: which is a mark on the vertex during the search algorithm. It is of String data type and should be set to "white" by default.
- `d`: which is an Integer denoting the distance from other Vertex to the current Vertex in Breath-First-Search. This is also used to record discovery time in Depth-First-Search. This attribute should be initialized to `sys.maxsize` at the start.
- `f`: which is an Integer denoting the final time in Depth-First-Search. This attribute should be initialized to `sys.maxsize` at the start.
- `parent`: which is a reference to the parent Vertex object. This attribute should be set to `None` at the start.

```python
import sys

class VertexSearch(Vertex):
    def __init__(self, id_=""):
        super().__init__(id_)
        self.colour: str = "white"
        self.d: int = sys.maxsize
        self.f: int = sys.maxsize
        self.parent: VertexSearch | None = None
    


```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
import sys

v = VertexSearch()
assert v.id_ == ""
assert v.colour == "white"
assert v.d == sys.maxsize
assert v.f == sys.maxsize
assert v.parent == None

parent_method = getattr(v, 'get_neighbours', None)
assert callable(parent_method)
parent_method = getattr(v, 'get_weight', None)
assert callable(parent_method)

```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5.** **You should do this after you completed HW2.** Create a class `Search2D` which takes in an object `GraphSearch` for its initialization. The class should have the following methods:
- `clear_vertices()`: which sets the attributes f all the vertices:
  - `colour` to "white"
  - `d` to `sys.maxsize`
  - `f` to `sys.maxsize`
  - `parent` to `None`.
 




```python
#Copy over GraphSearch from Homework
class GraphSearch(Graph):
    def _create_vertex(self, id_):
        return VertexSearch(id_)
```

```python
import sys

class Search2D:
    def __init__(self, g: Graph):
        self.graph = g
    
    def clear_vertices(self):
        for vertex in self.graph.vertices.values():
            vertex.colour = "white"
            vertex.d = sys.maxsize
            vertex.f = sys.maxsize
            vertex.parent = None
    
    def __iter__(self):
        return iter([v for v in self.graph])
    
    def __len__(self):
        return len([v for v in self.graph.vertices])
```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
g4 = GraphSearch()
g4.add_vertex("A")
g4.add_vertex("B")
g4.add_vertex("C")
g4.add_vertex("D")
g4.add_vertex("E")
g4.add_vertex("F")
g4.add_edge("A", "B")
g4.add_edge("A", "C")
g4.add_edge("B", "C")
g4.add_edge("B", "D")
g4.add_edge("C", "D")
g4.add_edge("D", "C")
g4.add_edge("E", "F")
g4.add_edge("F", "C")
gs4 = Search2D(g4)
gs4.clear_vertices()

assert len(gs4) == 6
assert [v.id_ for v in gs4] == ["A", "B", "C", "D", "E", "F"]
assert [v.colour for v in gs4] == ["white" for v in range(len(gs4))]
assert [v.d for v in gs4] == [sys.maxsize for v in range(len(gs4))]
assert [v.f for v in gs4] == [sys.maxsize for v in range(len(gs4))]
assert [v.parent for v in gs4] == [None for v in range(len(gs4))]
```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS6.** Create a class `SearchBFS` which is a subclass of `Search2D`. This subclass should implement the Breadth First Search algorithm in the following methods:

- `search_from(start)`: which initializes the `d` and `parent` attributes of each vertices in the graph from the `start` Vertex following Breadth-First-Search algorithm. Use your previous code that implements `Queue` data structure. 
- `get_shortest_path(start, dest)`: which returns a list of vertex ids that forms a shortest path from Vertex `start` to Vertex `dest`. This method should call `get_path()` (see next method in the list) and pass on an empty list as one of the input arguments. The method `get_path()` will populate this list if there is a path. 
    - If the path list is empty after calling `get_path()`, this means that either the starting vertex or the destination vertex do not exist in the grapth. In this case, exit the function returning a `None` object.
    - If the path list is not empty, it will either contain `No Path` as one of the items or a list of vertices that gives the path from the starting vertex to the destination vertex. In this case, simply return the list as it is. 
- `get_path(start, dest, result)`: which modifies the input list `result`. 
    - This method should first check whether the starting vertex and the destination vertex exist in the grapth. If they do not exist in either case, the method should exit returning a `None` object. 
    - If the starting and destination vertex exist in the graph, this method should call `search_from()` when the distance at `start` Vertex is not zero. A non-zero distance at the starting vertex means that we have not run the BFS algorithm from that starting vertex.
    - if the destination and the starting vertex are the same, modify the `result` list with this one vertex. This means that we have found the path that consists of only one vertex. 
    - if the destination vertex has no parent, this means there is no path. Add `No Path` string into the `result` list. 
    - otherwise, recursively call `get_path()` and add the result into the `result` list. 

```python
class Stack:
    def __init__(self):
        self._items = []
        
    def push(self, item):
        self._items.append(item)

    def pop(self):
        if self.is_empty:
            return None
        return self._items.pop()

    def peek(self):
        if self.is_empty:
            return None
        return self._items[-1]

    @property
    def is_empty(self):
        return self.size == 0

    @property
    def size(self):
        return len(self._items)

class Queue:
    def __init__(self):
        self.left_stack = Stack()
        self.right_stack = Stack()
    
    def enqueue(self, new_item):
        self.right_stack.push(new_item)

    def dequeue(self):
        if self.left_stack.is_empty:
            self.left_stack._items = self.right_stack._items[::-1]
            self.right_stack._items.clear()
        return self.left_stack.pop()

    def peek(self):
        if self.is_empty:
            return None
        elif self.left_stack.is_empty:
            return self.right_stack._items[0]
        else:
            return self.left_stack.peek()
    
    @property
    def is_empty(self):
        return self.left_stack.is_empty and self.right_stack.is_empty
    
    @property
    def size(self):
        return self.left_stack.size + self.right_stack.size
        

```

```python
class SearchBFS(Search2D):

    def search_from(self, start: str):
        self.clear_vertices()
        search_queue = Queue()
        start_vertex = self.graph.get_vertex(start)
        start_vertex.d = 0
        start_vertex.colour = 'black'
        search_queue.enqueue(start_vertex)
        while not search_queue.is_empty:
            current_vertex: VertexSearch = search_queue.dequeue()
            print(f'current_vertex = {current_vertex}')
            for neighbour in current_vertex.neighbours:
                if neighbour.colour == 'black': continue
                
                neighbour.colour = 'black'
                neighbour.parent = current_vertex
                neighbour.d = current_vertex.d + 1
                search_queue.enqueue(neighbour)
    
    def get_shortest_path(self, start, dest):
        res = []
        self.get_path(start, dest, res)
        if res is None or len(res) == 0:
            return None
        return res
    
    def get_path(self, start: str, dest: str, result: list) -> list | None:
        if self.graph.get_vertex(start) is None or self.graph.get_vertex(dest) is None:
            return None
        if len(result) == 0:
            self.search_from(start)
            result.append(dest)
        
        start_vertex: VertexSearch = self.graph.get_vertex(start)
        end_vertex: VertexSearch = self.graph.get_vertex(dest)

        if start_vertex == end_vertex:
            # we're done!
            return
        elif end_vertex.parent is None:
            result.insert(0, "No Path")
        else:
            result.insert(0, end_vertex.parent.id_)
            self.get_path(start, end_vertex.parent.id_, result)

        
```

```python

```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
g4 = GraphSearch()
g4.add_vertex("A")
g4.add_vertex("B")
g4.add_vertex("C")
g4.add_vertex("D")
g4.add_vertex("E")
g4.add_vertex("F")
g4.add_edge("A", "B")
g4.add_edge("A", "C")
g4.add_edge("B", "C")
g4.add_edge("B", "D")
g4.add_edge("C", "D")
g4.add_edge("D", "C")
g4.add_edge("E", "F")
g4.add_edge("F", "C")
gs4 = SearchBFS(g4)

gs4.search_from("A")
assert gs4.graph.get_vertex("A").d == 0
assert gs4.graph.get_vertex("A").colour == "black"
assert gs4.graph.get_vertex("A").parent == None
assert gs4.graph.get_vertex("B").d == 1
assert gs4.graph.get_vertex("B").colour == "black"
assert gs4.graph.get_vertex("B").parent == gs4.graph.get_vertex("A")
assert gs4.graph.get_vertex("C").d == 1
assert gs4.graph.get_vertex("C").colour == "black"
assert gs4.graph.get_vertex("C").parent == gs4.graph.get_vertex("A")
assert gs4.graph.get_vertex("D").d == 2
assert gs4.graph.get_vertex("D").colour == "black"
gs4.graph.get_vertex("D").parent
#assert gs4.graph.get_vertex("D").parent == gs4.graph.get_vertex("B")
ans = gs4.get_shortest_path("A", "D")
print(ans)
assert ans == ["A", "B", "D"]
ans = gs4.get_shortest_path("E", "D")
assert ans == ["E", "F", "C", "D"]
```

Output:
```
current_vertex = Vertex A is connected to: B, C
current_vertex = Vertex B is connected to: C, D
current_vertex = Vertex C is connected to: D
current_vertex = Vertex D is connected to: C
current_vertex = Vertex A is connected to: B, C
current_vertex = Vertex B is connected to: C, D
current_vertex = Vertex C is connected to: D
current_vertex = Vertex D is connected to: C
['A', 'B', 'D']
current_vertex = Vertex E is connected to: F
current_vertex = Vertex F is connected to: C
current_vertex = Vertex C is connected to: D
current_vertex = Vertex D is connected to: C
```

<!-- grade_id: cs62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```