# Week 2 Problem Set

## Cohort Sessions

**CS1.** *Binary Heap:* Write the following functions:



- `parent_of(index)`: returns the index of node's parent
- `left_of(index)`: returns the index of node's left child
- `right_of(index)`: returns the index of node's right child
- `max_child(array, index, heap_size)`: returns the index of node's largest child. You can assume that the node has at least one child.

Hint: 
- `index` starts from 0. 
- You can refer to the pseudocode in [Binary Heap and Heapsort](https://data-driven-world.github.io/2023/notes/category/week-2-analysing-programs).
- When finding the index of the largest child, consider the following cases:
    1. when the node only has one child
    2. when the node has two children


```python
def parent_of(index):
    return (index + 1) / 2
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert parent_of(1) == 0 
assert parent_of(2) == 0
assert parent_of(5) == 2
assert parent_of(6) == 2
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def left_of(index):
    return (2 * index) + 1
```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
assert left_of(0) == 1
assert left_of(1) == 3
assert left_of(3) == 7
assert left_of(6) == 13
```

<!-- grade_id: cs14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# technically correct!
# in a complete tree, the index of the right child is *indeed* the index of the left child + 1
def right_of(index):
    return left_of(index) + 1
```

<!-- grade_id: cs15, points: 1, grade: true, solution: false -->
```python
assert right_of(0) == 2
assert right_of(1) == 4
assert right_of(3) == 8
assert right_of(5) == 12
```

<!-- grade_id: cs16, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Re-order the code for returning the index of the max child from a given node index and some particular heap size.

Input:
- array: containing the binary heap data
- index: node index which we want to find the max child
- heap_size: the number of elements in the heap

Note that the array size can be bigger than the heap size. In this case, ignore the rest of the element as NOT part of the heap.

Output:
- index of the max child

```python
codes = {
    'e': 'def max_child(array, index, heap_size):',
        'c': 'if right_of(index) >= heap_size:',
            'b': 'return left_of(index)',
        'd': 'else:',
            'f': 'if array[left_of(index)] > array[right_of(index)]:',
                'b': 'return left_of(index)',
            'd': 'else:',
                'a': 'return right_of(index)',
}

```

```python
# Fill in the sequence and add "sub" or "exit sub" as needed.
answers = ['e', 'sub', 'c', 'sub', 'b', 'exit sub', 'd', 'sub', 'f', 'sub', 'b', 'exit sub', 'd', 'sub', 'a']
###
### YOUR CODE HERE
###

```

<!-- grade_id: cs17, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

Write the code below and test.

```python
def max_child(array, index, heap_size):
    if right_of(index) >= heap_size:
        return left_of(index)
    else:
        if array[left_of(index)] > array[right_of(index)]:
            return left_of(index)
        else:
            return right_of(index)
```

<!-- grade_id: cs18, points: 1, grade: true, solution: false -->
```python
maxheap = [16, 14, 10, 8, 7, 9, 3, 2, 4, 1] 
assert max_child(maxheap, 0, len(maxheap)) == 1
assert max_child(maxheap, 2, len(maxheap)) == 5
assert max_child(maxheap, 3, len(maxheap)) == 8
assert max_child(maxheap, 1, len(maxheap)) == 3
assert max_child(maxheap, 4, len(maxheap)) == 9
```

<!-- grade_id: cs19, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Binary Heap:* Write two functions 

- `max_heapify(array, index, size)`: that moves the node down so as to satisfy the heap property.  The first argument is the array that contains the heap. The second argument is an integer index where to start the process of heapifying. The third argument is the size of the heap in the array. This argument will be useful in heapsort algorithm where we take out the elements in the array from the heap. 

Hint: You should make use of `size` argument to determine the last element of the heap in the array rather than `len(array)`.

- `build_max_heap(heap)`: that builds the max heap from any array. This function should make use of `max_heapify()` in its definition.

Hint: You can refer to the pseudocode in [Binary Heap and Heapsort](https://data-driven-world.github.io/2023/notes/category/week-2-analysing-programs) for the above functions.

```python
def max_heapify(array, index, size):
    left, right = left_of(index), right_of(index)
    is_leaf = left >= size
    is_left_only = right >=  size
    if is_leaf or (array[index] >= array[left] and (is_left_only or array[index] >= array[right])):
        return
    
    if is_left_only or array[left] > array[right]:
        array[left], array[index] = array[index], array[left]
        max_heapify(array, left, size)
    else:
        array[right], array[index] = array[index], array[right]
        max_heapify(array, right, size)
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def build_max_heap(array):
    for i in range(len(array) // 2 - 1, -1, -1):
        max_heapify(array, i, len(array))

```

<!-- grade_id: cs23, points: 1, grade: true, solution: false -->
```python
array = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]

build_max_heap(array)
print(array)
assert array == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

<!-- grade_id: cs24, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Heapsort:* Implement heapsort algorithm following the pseudocode in [Binary Heap and Heapsort](https://data-driven-world.github.io/2023/notes/category/week-2-analysing-programs).



```python
def heapsort(array):
    heap = array[:]
    n_sorted_items = 0
    while len(heap) > 0:
        build_max_heap(heap)
        n_sorted_items += 1
        array[-n_sorted_items] = heap[0]
        heap[0] = heap[-1]
        del heap[-1]
    
    
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
array = [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
heapsort(array)
print(array)
assert array == [1, 2, 3, 4, 7, 8, 9, 10, 14, 16]
```

Output:
```
[1, 2, 3, 4, 7, 8, 9, 10, 14, 16]
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Measure computational time of Python's built-in sort function by filling the template below. Hint: 
- You will need the function `gen_random_int()` from Week 01 Problem Set.
- Use `sorted(list)` function of Python's list See [Python's Sorting HOW TO Documentation](https://docs.python.org/3.3/howto/sorting.html). 

```python
import time
import random

def run_function(f, *x):
    start = time.time()
    f(*x)
    end = time.time()
    return end-start

def gen_random_int(number, seed):
    random.seed(seed)
    return random.sample(range(0, number), k=number)

 
time_builtin = []
# set the maximum power for 10^power number of inputs
maxpower = 7

for n in range(1, maxpower + 1):
    # create array for 10^1, 10^2, etc 
    # use seed 100
    array = gen_random_int(10 ** n, 100)
    
    # call run_function with sorted() and array as arguments
    
    result = run_function(sorted, array) # 3.044 seconds
    # result = run_function(gen_random_int, 10 ** n, 100) # 5.00 seconds
    
    time_builtin.append(result)


print(time_builtin)
# print(sum(time_builtin))
```

Output:
```
[0.0, 0.0, 0.0, 0.0010004043579101562, 0.01372528076171875, 0.21276187896728516, 2.867779016494751]
3.095266580581665
```